---
  title: "Introduction to R"
  description: "Today we are going to learn the basics of R. "
---

## R Addition

```yaml
type: NormalExercise 
xp: 100 
key: 9ed7531164   
```


We will start off seeing how R can be used as a calculator.


`@instructions`
Add together 7 and 11

`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}
#### To add 3 and 5 type
3 + 5

### Now add 7 and 11
```

`@solution`

```{r}
#### To add 3 and 5 type
3 + 5

### Now add 7 and 11 
7 + 11
```

`@sct`

```{r}
ex() %>% check_output_expr("7+11")
```

---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: c02ebefc9d   
```


R can do more than just addition but can also do subtraction, multiplication and division.


`@instructions`
Do the following problems:

Subtract 3 from 8. 
Multiply 2 by 6
Divide 9 by 2

`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}
### subtraction uses -

## multiplication uses *

## division uses / 
```

`@solution`

```{r}
### subtraction uses -
8 - 3

## multiplication uses *
6 * 2

## division uses / 
9 / 2
```

`@sct`

```{r}

```

